rule_files:
  - ../prometheus_alert_rules.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: 'NooBaa_system_capacity'
        values: '86+0x16'
    alert_rule_test:
      - alertname: 'NooBaaSystemCapacityWarning85'
        eval_time: 11m
        exp_alerts:
         - exp_labels:
             severity: warning
           exp_annotations:
             message: 'A NooBaa System Is Approaching Its Capacity'
             description: 'A NooBaa system is approaching its capacity, usage is more than 85%'
             storage_type: 'NooBaa'
             severity_level: 'warning'

  - interval: 1m
    input_series:
      - series: 'NooBaa_system_capacity'
        values: '96+0x16'
    alert_rule_test:
      - alertname: 'NooBaaSystemCapacityWarning95'
        eval_time: 11m
        exp_alerts:
         - exp_labels:
             severity: critical
           exp_annotations:
             message: 'A NooBaa System Is Approaching Its Capacity'
             description: 'A NooBaa system is approaching its capacity, usage is more than 95%'
             storage_type: 'NooBaa'
             severity_level: 'critical'

  - interval: 1m
    input_series:
      - series: 'NooBaa_system_capacity'
        values: '100+0x16'
    alert_rule_test:
      - alertname: 'NooBaaSystemCapacityWarning100'
        eval_time: 11m
        exp_alerts:
         - exp_labels:
             severity: critical
           exp_annotations:
             message: 'A NooBaa System Approached Its Capacity'
             description: 'A NooBaa system approached its capacity, usage is at 100%'
             storage_type: 'NooBaa'
             severity_level: 'critical'
             
  - interval: 1m
    input_series:
      - series: 'NooBaa_resource_status{resource_name="pool1"}'
        values: '0+0x16'
    alert_rule_test:
      - alertname: 'NooBaaResourceErrorState'
        eval_time: 11m
        exp_alerts:
         - exp_labels:
             severity: warning
             resource_name: pool1
           exp_annotations:
             message: 'A NooBaa Resource Is In Error State'
             description: 'A NooBaa resource pool1 is in error state for more than 6m'
             storage_type: 'NooBaa'
             severity_level: 'warning'

  - interval: 1m
    input_series:
      - series: 'NooBaa_bucket_status{bucket_name="bucket1"}'
        values: '0+0x16'
    alert_rule_test:
      - alertname: 'NooBaaBucketErrorState'
        eval_time: 11m
        exp_alerts:
         - exp_labels:
             severity: warning
             bucket_name: bucket1
           exp_annotations:
             message: 'A NooBaa Bucket Is In Error State'
             description: 'A NooBaa bucket bucket1 is in error state for more than 6m'
             storage_type: 'NooBaa'
             severity_level: 'warning'

  - interval: 1m
    input_series:
      - series: 'NooBaa_bucket_quota{bucket_name="bucket1"}'
        values: '81+0x16'
    alert_rule_test:
      - alertname: 'NooBaaBucketReachingQuotaState'
        eval_time: 11m
        exp_alerts:
         - exp_labels:
             severity: warning
             bucket_name: bucket1
           exp_annotations:
             message: 'A NooBaa Bucket Is In Reaching Quota State'
             description: 'A NooBaa bucket bucket1 is using 81% of its quota'
             storage_type: 'NooBaa'
             severity_level: 'warning'

  - interval: 1m
    input_series:
      - series: 'NooBaa_bucket_quota{bucket_name="bucket1"}'
        values: '102+0x16'
    alert_rule_test:
      - alertname: 'NooBaaBucketExceedingQuotaState'
        eval_time: 11m
        exp_alerts:
         - exp_labels:
             severity: warning
             bucket_name: bucket1
           exp_annotations:
             message: 'A NooBaa Bucket Is In Exceeding Quota State'
             description: 'A NooBaa bucket bucket1 is exceeding its quota - 102% used'
             storage_type: 'NooBaa'
             severity_level: 'warning'
  
  - interval: 1m
    input_series:
      - series: 'NooBaa_bucket_capacity{bucket_name="bucket1"}'
        values: '81+0x16'
    alert_rule_test:
      - alertname: 'NooBaaBucketLowCapacityState'
        eval_time: 11m
        exp_alerts:
         - exp_labels:
             severity: warning
             bucket_name: bucket1
           exp_annotations:
             message: 'A NooBaa Bucket Is In Low Capacity State'
             description: 'A NooBaa bucket bucket1 is using 81% of its capacity'
             storage_type: 'NooBaa'
             severity_level: 'warning'

  - interval: 1m
    input_series:
      - series: 'NooBaa_bucket_capacity{bucket_name="bucket1"}'
        values: '100+0x16'
    alert_rule_test:
      - alertname: 'NooBaaBucketNoCapacityState'
        eval_time: 11m
        exp_alerts:
         - exp_labels:
             severity: warning
             bucket_name: bucket1
           exp_annotations:
             message: 'A NooBaa Bucket Is In No Capacity State'
             description: 'A NooBaa bucket bucket1 is using all of its capacity'
             storage_type: 'NooBaa'
             severity_level: 'warning'
